```python
pip install nose
```

Output:
```
Collecting nose
  Downloading nose-1.3.7-py3-none-any.whl (154 kB)
[2K     [90m━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━━[0m [32m154.7/154.7 kB[0m [31m1.1 MB/s[0m eta [36m0:00:00[0m00:01[0m00:01[0m
[?25hInstalling collected packages: nose
Successfully installed nose-1.3.7
Note: you may need to restart the kernel to use updated packages.
```

```python
#Perform logistic regression on diabetes dataset
```

<!-- grade_id: cell-65901a02d7b3d454, points: 0, grade: false, solution: true -->
```python
#Load diabetes dataset here and return shape of dataset
from cz_utility import dataset
# diabetes dataset
df = dataset.load('63ea33b29ac5f900139b8d07')
def returnShape(df):
    # your code here
    return df.shape
    raise NotImplementedError
```

<!-- grade_id: cell-9d4cd5b17d896c66, points: 2, grade: true, solution: false -->
```python
from nose.tools import assert_equal
assert_equal(returnShape(df), (768, 9))
```

```python
df.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 768 entries, 0 to 767
Data columns (total 9 columns):
 #   Column                    Non-Null Count  Dtype  
---  ------                    --------------  -----  
 0   Pregnancies               768 non-null    int64  
 1   Glucose                   768 non-null    int64  
 2   BloodPressure             768 non-null    int64  
 3   SkinThickness             768 non-null    int64  
 4   Insulin                   768 non-null    int64  
 5   BMI                       768 non-null    float64
 6   DiabetesPedigreeFunction  768 non-null    float64
 7   Age                       768 non-null    int64  
 8   Outcome                   768 non-null    int64  
dtypes: float64(2), int64(7)
memory usage: 54.1 KB
```

<!-- grade_id: cell-97b7b8d9c24aa579, points: 0, grade: false, solution: true -->
```python
#separate input and output from dataset and remove unwanted data columns
#take x for input and y for output
def sepAndRemove(df):
    # your code here
    y = df['Outcome'].copy()
    x = df.drop(['Outcome','SkinThickness'], axis=1).copy()
    return [x, y]
    raise NotImplementedError
```

<!-- grade_id: cell-f7e39f11fcd55f39, points: 2, grade: true, solution: false -->
```python
trainAndTestData = sepAndRemove(df)
assert_equal(trainAndTestData[0].shape, (768, 7))
assert_equal(trainAndTestData[1].shape, (768,))
```

<!-- grade_id: cell-05c4eb3e485c557a, points: 0, grade: false, solution: true -->
```python
#separate training and testing dataset with test_size = 0.3
from sklearn.model_selection import train_test_split
def trainTestSplit(trainAndTestData):
    # your code here
    X_train, X_test, Y_train, Y_test = train_test_split(trainAndTestData[0], trainAndTestData[1], test_size=0.30, random_state=0)
    return [X_train, X_test, Y_train, Y_test]
    raise NotImplementedError
```

<!-- grade_id: cell-a7a6e742dbef0d20, points: 2, grade: true, solution: false -->
```python
trainAndTestData=trainTestSplit(sepAndRemove(df))
assert_equal(trainAndTestData[0].shape,(537, 7))
assert_equal(trainAndTestData[1].shape,(231, 7))
assert_equal(trainAndTestData[2].shape,(537,))
assert_equal(trainAndTestData[3].shape,(231,))
```

<!-- grade_id: cell-6e973aeb54f8a469, points: 0, grade: false, solution: true -->
```python
#use MinMaxScaler to scale the data to a fixed range
from sklearn.preprocessing import MinMaxScaler
def scalar(trainAndTestData):
    # your code here
    scaler = MinMaxScaler()
    trainAndTestData[0] = scaler.fit_transform(trainAndTestData[0])
    trainAndTestData[1] = scaler.fit_transform(trainAndTestData[1])
    return trainAndTestData
    raise NotImplementedError
```

<!-- grade_id: cell-0bb62bf101293688, points: 2, grade: true, solution: false -->
```python
trainAndTestData=scalar(trainTestSplit(sepAndRemove(df)))
assert_equal(trainAndTestData[0].shape, (537, 7))
assert_equal(trainAndTestData[1].shape, (231, 7))

```

<!-- grade_id: cell-a9178e0adf01003e, points: 0, grade: false, solution: true -->
```python
#perform logistic regression here and return accuracy between predected output for testx and testy
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import accuracy_score
def logisticReg(trainAndTestData, trainy, testy):
    # your code here
    model = LogisticRegression()
    model.fit(trainAndTestData[0], trainy)
    predictions = model.predict(trainAndTestData[1])
    score = accuracy_score(testy, predictions)
    return score
    raise NotImplementedError
```

<!-- grade_id: cell-b26830e8ae1d3a06, points: 2, grade: true, solution: false -->
```python
#there might be difference between student answer and faculty answer
trainAndTestData1 = trainTestSplit(sepAndRemove(df))
accuracyScore = logisticReg(scalar(trainTestSplit(sepAndRemove(df))), trainAndTestData1[2], trainAndTestData1[3])
assert_equal(accuracyScore>0.60, True)
```

```python

```